# Einführung

Jeder Abschnitt wird von einer größeren Übungsaufgabe begleitet. Während in den Abschnitten 1-3 die Übungen mit Jupyter erledigt werden könne, wird in den Abschnitten 4 & 5 mit Android Studio gearbeitet (näheres dazu folgt in Abschnitt 4). In den ersten 3 Übungen wird ein seperates Notebook angeboten, das Sie runterladen und bearbeiten sollen. Eine Bearbeitung der Übungsaufgaben ist für das Abschlusszertifikat nötig. In der Aufgabenbeschreibung wird angegeben, wie die zu implementierenden Methoden oder Klassen benannt werden sollen. Bitte halten Sie sich an diese Vorgaben, da nur so die Tests, mit denen Sie Ihren Code ausprobieren können, funktionieren.


# Aufgabe - Passwörter

Heutzutage ist es immer wichtiger sichere Passwörter zu benutzen. Ob es der Zugang zum Online-Banking oder das Passwort für Netflix ist. Ein fremder Zugang kann einen großen Schaden verursachen. Deswegen soll im Folgenden ein Passwortgenerator und -validator programmiert werden, der sichere Passwörter erkennt und auch vorschlägt.<br />
Ein Passwort ist in dieser Aufgabe als "sicher" anzusehen, wenn es die [Standards des BSI](https://www.bsi.bund.de/DE/Themen/Verbraucherinnen-und-Verbraucher/Informationen-und-Empfehlungen/Cyber-Sicherheitsempfehlungen/Accountschutz/Sichere-Passwoerter-erstellen/sichere-passwoerter-erstellen_node.htmlhttps://www.bsi.bund.de/DE/Themen/Verbraucherinnen-und-Verbraucher/Informationen-und-Empfehlungen/Cyber-Sicherheitsempfehlungen/Accountschutz/Sichere-Passwoerter-erstellen/sichere-passwoerter-erstellen_node.html) erfüllt:
- 20 bis 25 Zeichen lang und zwei Zeichenarten genutzt (beispielsweise eine Folge von Wörtern).
- 8 bis 12 Zeichen lang und vier Zeichenarten genutzt. 
- 8 Zeichen lang, drei Zeichenarten genutzt und es zusätzlich durch eine Mehr-Faktor-Authentisierung abgesichert (beispielsweise durch einen Fingerabdruck, eine Bestätigung per App oder eine PIN).

Der dritten Stichpunkt soll im Folgenden ignoriert werden, da eine Mehr-Faktor-Authentisierung die Aufgabe übersteigen würde. Es soll sich auf die ersten beiden Stichpunkte fokussiert werden.

Die 4 Zeichenarten sind:
- Großbuchstaben
- Kleinbuchstaben
- Zahlen
- Sonderzeichen (zur Vereinfachnung nur `@`, `/`, `#`, `+`, `-`, `*`, `!`, `?`, `&`)

## Anzahl der Zeichenarten
Implementieren Sie als Vorbereitung auf die eigentliche Aufgabe Erweiterungsmethoden für die Klasse `String`.
- `anzahlKleinbuchstaben`: gibt die Anzahl an Kleinbuchstaben zurück
- `anzahlGrossbuchstaben`: gibt die Anzahl an Großbuchstaben zurück
- `anzahlZahlen`: gibt die Anzahl an Zahlen zurück
- `anzahlSonderzeichen`: gibt die Anzahl der Sonderzeichen zurück
- `anzahlZeichenarten`: gibt die Anzahl der Zeichenarten zurück

<!-- grade_id: abschnitt1-1-code, points: 0, grade: false, solution: true -->
```kotlin
fun String.anzahlZeichenarten(): Int{
    var counter = 0
    if (this.anzahlGrossbuchstaben() > 0) 
        counter++
    if (this.anzahlKleinbuchstaben() > 0) 
        counter++
    if (this.anzahlZahlen() > 0) 
        counter++
    if (this.anzahlSonderzeichen() > 0) 
        counter++
    return counter
}

fun String.anzahlKleinbuchstaben(): Int{
    var counter = 0
    for(char in this){
        if (char in 'a'..'z'){
            counter++
        }
    }
    return counter
}

fun String.anzahlGrossbuchstaben(): Int{
    var counter = 0
    for(char in this){
        if (char in 'A'..'Z'){
            counter++
        }
    }
    return counter
}

fun String.anzahlZahlen(): Int{
    var counter = 0
    for(char in this){
        if (char in '0'..'9'){
            counter++
        }
    }
    return counter
}

fun String.anzahlSonderzeichen(): Int{
    val sonderzeichen = listOf('@', '/', '#', '+', '-', '*', '!', '?', '&')
    var counter = 0
    for(char in this){
        if (char in sonderzeichen){
            counter++
        }
    }
    return counter
}
```

<!-- grade_id: abschnitt1-1-test1, points: 10, grade: true, solution: false -->
```kotlin
//Test von anzahlZeichenarten
val pwTests = mapOf("" to 0, "aaaaa" to 1, "a2" to 2, "D#" to 2, "#1" to 2, "aB+" to 3, "aaBDA2834" to 3, "adhGDUWahduWe" to 2, "234#34#123+34Hca" to 4)
for((pw, anzahl) in pwTests){
    if(pw.anzahlZeichenarten() != anzahl){
        throw AssertionError("wrong")
    }
}
```

<!-- grade_id: abschnitt1-1-test2, points: 0, grade: true, solution: false -->
```kotlin
//Test von anzahlZeichenarten
val pwTests = mapOf("" to 0, "aaaaa" to 1, "a2" to 2, "D#" to 2, "#1" to 2, "aB+" to 3, "aaBDA2834" to 3, "adhGDUWahduWe" to 2, "234#34#123+34Hca" to 4)
for((pw, anzahl) in pwTests){
    assert(pw.anzahlZeichenarten() == anzahl)
}
```

## Passwortstärke
Nun soll die Stärke des Passworts von 0 - 100 berechnet werden. Die Berechnung ist an [Passwordmeter](http://www.passwordmeter.com/) angelehnt. Folgende Attribute geben Punkte:

<table style="font-size: 18px; width: 100%">
<thead>
  <tr>
    <th>Eigenschaft</th>
    <th style="width: 20%; min-width: 100px">Punkte</th>
  </tr>
</thead>
<tbody>
  <tr>
    <td>Anzahl Zeichen insgesamt</td>
    <td>+(n*4)</td>
  </tr>
  <tr>
    <td>Anzahl Großbuchstaben</td>
    <td>+((len-n)*2)</td>
  </tr>
  <tr>
    <td>Anzahl Kleinbuchstaben</td>
    <td>+((len-n)*2)</td>
  </tr>
  <tr>
    <td>Anzahl Zahlen</td>
    <td>+(n*4)</td>
  </tr>
  <tr>
    <td>Anzahl Sonderzeichen</td>
    <td>+(n*6)</td>
  </tr>
  <tr>
    <td>Anzahl Zeichenarten (2-4)</td>
    <td>+(n*3)</td>
  </tr>
  <tr>
    <td>Nur 1 Zeichenart</td>
    <td>-len</td>
  </tr>
  <tr>
      <td>Wiederholende Zeichen (Bsp: "<b>aaaa</b>" n=4, "<b>a</b>C<b>w</b>i3Wf<b>a</b>sd<b>aw</b>" n=5)</td>
    <td>-(n*3)</td>
  </tr>
</tbody>
</table>

Implementieren Sie eine Methode `passwortstaerke`, die das Passwort (als `String`) übergeben bekommt. Zurückgegeben soll die Stärke des Passworts, welche mit der obigen Tabelle berechnet werden soll. Der beste Wert ist 100, der Schlechteste 0.

<!-- grade_id: abschnitt1-2-code, points: 0, grade: false, solution: true -->
```kotlin

fun passwortstaerke(pw: String) : Int{
    return 1
}

```

<!-- grade_id: abschnitt1-2-test1, points: 10, grade: true, solution: false -->
```kotlin
//Test von anzahlZeichenarten
val pwTests = mapOf("" to 0, "aaaaa" to 3, "a2" to 20, "D#" to 22, "jdfVqed&213" to 98, "AAAaaa111" to 54, "aaBDA2834" to 81, "adhGDUWahduWe" to 60, "234#34#123+34Hca" to 100)
for((pw, staerke) in pwTests){
    if(passwortstaerke(pw) != staerke){
        throw AssertionError("wrong")
    }
}
```

Output:
```
wrong
java.lang.AssertionError: wrong
org.jetbrains.kotlinx.jupyter.ReplEvalRuntimeException: wrong
	at org.jetbrains.kotlinx.jupyter.repl.impl.InternalEvaluatorImpl.eval(InternalEvaluatorImpl.kt:91)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:63)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withHost(repl.kt:536)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl.execute(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.repl.CellExecutor$DefaultImpls.execute$default(CellExecutor.kt:13)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:371)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:361)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withEvalContext(repl.kt:346)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.eval(repl.kt:3
```

<!-- grade_id: abschnitt1-2-test2, points: 10, grade: true, solution: false -->
```kotlin
//Test von anzahlZeichenarten
val pwTests = mapOf("" to 0, "aaaaa" to 3, "a2" to 20, "D#" to 22, "jdfVqed&213" to 98, "AAAaaa111" to 54, "aaBDA2834" to 81, "adhGDUWahduWe" to 60, "234#34#123+34Hca" to 100)
for((pw, staerke) in pwTests){
    assert(passwortstaerke(pw) == staerke)
}
```

## Passwortgenerator

Unter Verwendung der `passwordstaerke`-Methode soll nun mit Hilfe der [Brute-Force-Methode](https://de.wikipedia.org/wiki/Brute-Force-Methode) ein Passwort generiert werden. Implementieren Sie dazu eine Methode `passwortGenerieren`. Die Stärke des generierten Passworts soll eine bestimmte Stärke, die übergeben werde soll, übersteigen.

<!-- grade_id: abschnitt1-3-code, points: 0, grade: false, solution: true -->
```kotlin
fun passwortGenerieren(staerke : Int) : String{
    return "q"
}

```

<!-- grade_id: abschnitt1-3-test1, points: 10, grade: true, solution: false -->
```kotlin
//Test von passwortGenerieren
for (i in 0..100 step 5){
    val testPasswort = passwortGenerieren(i)
    if (passwortstaerke(testPasswort) < i){
        throw AssertionError()
    }
}
for (i in 0..5){
    val testPasswort = passwortGenerieren(100)
    if (passwortstaerke(testPasswort) < 100){
        throw AssertionError()
    }
}
```

Output:
```
java.lang.AssertionError
org.jetbrains.kotlinx.jupyter.ReplEvalRuntimeException: 
	at org.jetbrains.kotlinx.jupyter.repl.impl.InternalEvaluatorImpl.eval(InternalEvaluatorImpl.kt:91)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:63)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withHost(repl.kt:536)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl.execute(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.repl.CellExecutor$DefaultImpls.execute$default(CellExecutor.kt:13)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:371)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:361)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withEvalContext(repl.kt:346)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.eval(repl.kt:361)
	at org.jetbra
```

<!-- grade_id: abschnitt1-3-test2, points: 10, grade: true, solution: false -->
```kotlin
//Test von passwortGenerieren
for (i in 0..100 step 5){
    val testPasswort = passwortGenerieren(i)
    assert(passwortstaerke(testPasswort) >= i)
}
for (i in 0..5){
    val testPasswort = passwortGenerieren(100)
    assert(passwortstaerke(testPasswort) == 100)
}
```

<!-- grade_id: testtest, points: 1, grade: true, solution: false -->
```kotlin
throw AssertionError()
```

Output:
```
java.lang.AssertionError
org.jetbrains.kotlinx.jupyter.ReplEvalRuntimeException: 
	at org.jetbrains.kotlinx.jupyter.repl.impl.InternalEvaluatorImpl.eval(InternalEvaluatorImpl.kt:91)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:63)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withHost(repl.kt:536)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl.execute(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.repl.CellExecutor$DefaultImpls.execute$default(CellExecutor.kt:13)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:371)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:361)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withEvalContext(repl.kt:346)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.eval(repl.kt:361)
	at org.jetbra
```

Nun soll die Passwortgenerierung erweitert werden. Implementieren Sie eine zusätzliche Methode `passwortGenerierenErweitert`, die zusätzlich zu der Mindeststärke auch die Anzahl der Zeichen jeder Art übergeben werden. Standardmäßig soll die Mindeststärke 75 und die Anzahl an Zeichen jeder Art 0 sein. Die Zeichen sollen an zufälligen Stellen im Passwort zu finden sein. Ist nach 200 Versuchen kein Passwort für die Mindeststärke mit der Kombination an Zeichen möglich, soll ein Fehler geworfen/ausgegeben werden.<br />
Beispiel:<br />
`passwortGenerierenErweitert(80,3,4,2,1)`: Generiert ein Passwort, das mindestens die Stärke 80, 3 Großbuchstaben, 4 Kleinbuchstaben, 2 Zahlen und 1 Sonderzeichen besitzt.<br />
`passwortGenerierenErweitert(50,5,7)`: Generiert ein Passwort, das mindestens die Stärke 50, 5 Großbuchstaben, 7 Kleinbuchstaben, 0 Zahlen und 0 Sonderzeichen besitzt.<br />

<!-- grade_id: abschnitt1-4-code, points: 0, grade: false, solution: true -->
```kotlin
fun passwortGenerierenErweitert(staerke: Int = 75, vorGross: Int = 0, vorKlein: Int = 0, vorZahlen: Int = 0, vorSonder: Int = 0): String{
    // YOUR CODE HERE
    throw NotImplementedError()
}

```

<!-- grade_id: abschnitt1-4-test, points: 10, grade: true, solution: false -->
```kotlin
//Test von passwortGenerieren mit erweiterten Parametern
for(i in 0..49){
    val staerke = (0..100).random()
    val gross = (0..6).random()
    val klein = (0..6).random()
    val zahlen = (0..6).random()
    val sonder = (0..6).random()
    val testPasswort = passwortGenerierenErweitert(staerke, gross, klein, zahlen, sonder)
    if (
        (testPasswort.anzahlGrossbuchstaben() != gross || testPasswort.anzahlKleinbuchstaben() != klein ||
        testPasswort.anzahlZahlen() != zahlen || testPasswort.anzahlSonderzeichen() != sonder ||
        passwortstaerke(testPasswort) < staerke) && testPasswort != "Fehler"
    ){
        throw AssertionError()
    }
    
}
for(i in 0..4){
    val staerke = 100
    val gross = (0..6).random()
    val klein = (0..6).random()
    val zahlen = (0..6).random()
    val sonder = (0..6).random()
    val testPasswort = passwortGenerierenErweitert(staerke, gross, klein, zahlen, sonder)
    if (
        (testPasswort.anzahlGrossbuchstaben() != gross || testPasswort.anzahlKleinbuchstaben() != klein ||
        testPasswort.anzahlZahlen() != zahlen || testPasswort.anzahlSonderzeichen() != sonder ||
        passwortstaerke(testPasswort) < staerke) && testPasswort != "Fehler"
    ){
        throw AssertionError()
    }
}
val testPasswort = passwortGenerierenErweitert(10, 5, 2)
if (
    (testPasswort.anzahlGrossbuchstaben() != 5 || testPasswort.anzahlKleinbuchstaben() != 2 ||
    testPasswort.anzahlZahlen() != 0 || testPasswort.anzahlSonderzeichen() != 0 ||
    passwortstaerke(testPasswort) < 10) && testPasswort != "Fehler"
){
    throw AssertionError()
}
```

Output:
```
An operation is not implemented.
kotlin.NotImplementedError: An operation is not implemented.
org.jetbrains.kotlinx.jupyter.ReplEvalRuntimeException: An operation is not implemented.
	at org.jetbrains.kotlinx.jupyter.repl.impl.InternalEvaluatorImpl.eval(InternalEvaluatorImpl.kt:91)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:63)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl$execute$1$result$1.invoke(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withHost(repl.kt:536)
	at org.jetbrains.kotlinx.jupyter.repl.impl.CellExecutorImpl.execute(CellExecutorImpl.kt:62)
	at org.jetbrains.kotlinx.jupyter.repl.CellExecutor$DefaultImpls.execute$default(CellExecutor.kt:13)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:371)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl$eval$1.invoke(repl.kt:361)
	at org.jetbrains.kotlinx.jupyter.ReplForJupyterImpl.withEvalConte
```